Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name below:

```OCaml
NAME = ""
ROLLNO = ""
```

## Important notes about grading:

1. **Compiler errors:** All code you submit must compile. Programs that do not compile will probably receive an automatic zero. If you are having trouble getting your assignment to compile, please visit consulting hours. If you run out of time, it is better to comment out the parts that do not compile, than hand in a more complete file that does not compile.
2. **Late assignments:** Please carefully review the course website's policy on late assignments, as all assignments handed in after the deadline will be considered late. Verify on moodle that you have submitted the correct version, before the deadline. Submitting the incorrect version before the deadline and realizing that you have done so after the deadline will be counted as a late submission.

## Problem 1

Write a function 

```ocaml
cond_dup : 'a list -> ('a -> bool) -> 'a
```

list that takes in a list and preciate and duplicates all elements which satisfy the condition expressed in the predicate. For example, `cond_dup [3;4;5] (fun x -> x mod 2 = 1) = [3;3;4;5;5]`.

<!-- grade_id: cell-7950f1ea0a84b889, points: 0, grade: false, solution: true -->
```OCaml
let rec cond_dup l f =
  (* YOUR CODE HERE *)
   raise (Failure "Not implemented")
```

<!-- grade_id: cell-1d772b02ff64382c, points: 10, grade: true, solution: false -->
```OCaml
assert (cond_dup [3;4;5] (fun x -> x mod 2 = 1) = [3;3;4;5;5])
```

## Problem 2

Write a function 

```ocaml
n_times : ('a -> 'a) * int * 'a -> 'a
```

such that `n_times (f,n,v)` applies `f` to `v` `n` times. For example, `n_times((fun x-> x+1), 50, 0) = 50`. If `n<=0` return `v`.


<!-- grade_id: cell-c28f98bd63f56ba1, points: 0, grade: false, solution: true -->
```OCaml
let rec n_times (f, n, v) =
  (* YOUR CODE HERE *)
   raise (Failure "Not implemented")
```

<!-- grade_id: cell-98364653c4b77b2d, points: 10, grade: true, solution: false -->
```OCaml
assert (n_times((fun x-> x+1), 50, 0) = 50)
```

## Problem 3

Write a function 

```ocaml
range : int -> int -> int list 
```

such that `range num1 num2` returns an ordered list of all integers from `num1` to `num2` inclusive. For example, `range 2 5 = [2;3;4;5]`. Raise the exception `IncorrectRange` if `num2 < num1`.

<!-- grade_id: cell-1fd147f7eb3b139d, points: 0, grade: false, solution: true -->
```OCaml
exception IncorrectRange

let rec range num1 num2 =
  (* YOUR CODE HERE *)
   raise (Failure "Not implemented")
```

<!-- grade_id: cell-bba984120a807d42, points: 10, grade: true, solution: false -->
```OCaml
assert (range 2 5 = [2;3;4;5])
```

## Problem 4

Write a function 

```ocaml
zipwith : ('a -> 'b -> 'c) -> 'a list -> 'b list -> 'c list
```

such that `zipwith f l1 l2` generates a list whose `ith` element is obtained by applying `f` to the `ith` element of `l1` and the `ith` element of `l2` . If the lists have different lengths, the extra elements in the longer list should be ignored. For example, `zipwith (+) [1;2;3] [4;5] = [5;7]`.

<!-- grade_id: cell-932054b41dace011, points: 0, grade: false, solution: true -->
```OCaml
let rec zipwith f l1 l2 =
  (* YOUR CODE HERE *)
   raise (Failure "Not implemented")
```

<!-- grade_id: cell-55aefe5fa4f927d0, points: 10, grade: true, solution: false -->
```OCaml
assert (zipwith (+) [1;2;3] [4;5] = [5;7]) 
```

## Problem 5

Write a function 

```ocaml
buckets : ('a -> 'a -> bool) -> 'a list -> 'a list list 
```

that partitions a list into equivalence classes. That is, buckets equiv lst should return a list of lists where each sublist in the result contains equivalent elements, where two elements are considered equivalent if equiv returns true. For example:


```ocaml
buckets (=) [1;2;3;4] = [[1];[2];[3];[4]]
buckets (=) [1;2;3;4;2;3;4;3;4] = [[1];[2;2];[3;3;3];[4;4;4]]
buckets (fun x y -> (=) (x mod 3) (y mod 3)) [1;2;3;4;5;6] = [[1;4];[2;5];[3;6]]
```

The order of the buckets must reflect the order in which the elements appear in the original list. For example, the output of `buckets (=) [1;2;3;4]` should be `[[1];[2];[3];[4]]` and not `[[2];[1];[3];[4]]` or any other permutation.

The order of the elements in each bucket must reflect the order in which the elements appear in the original list. For example, the output of `buckets (fun x y -> (=) (x mod 3) (y mod 3)) [1;2;3;4;5;6]` should be `[[1;4];[2;5];[3;6]]` and not `[[4;1];[5;2];[3;6]]` or any other permutations.

Assume that the comparison function `('a -> 'a -> bool)` is commutative, associative and idempotent.

Just use lists. Do not use sets or hash tables.

List append function `@` may come in handy. `[1;2;3] @ [4;5;6] = [1;2;3;4;5;6]`.

<!-- grade_id: cell-f4d88f256cc9db8b, points: 0, grade: false, solution: true -->
```OCaml
let buckets p l =
  (* YOUR CODE HERE *)
   raise (Failure "Not implemented")
```

<!-- grade_id: cell-fdbd3b5564e151cb, points: 20, grade: true, solution: false -->
```OCaml
assert (buckets (=) [1;2;3;4] = [[1];[2];[3];[4]]);
assert (buckets (=) [1;2;3;4;2;3;4;3;4] = [[1];[2;2];[3;3;3];[4;4;4]]);
assert (buckets (fun x y -> (=) (x mod 3) (y mod 3)) [1;2;3;4;5;6] = [[1;4];[2;5];[3;6]])
```

## Problem 6

Write a function 

```ocaml
remove_stutter : 'a list -> 'a list
```

that removes stuttering from the original list. For example, `remove_stutter [1;2;2;3;1;1;1;4;4;2;2]= [1;2;3;1;4;2]`. Option type may come in handy. 

<!-- grade_id: cell-8865e0dd05cbd32d, points: 0, grade: false, solution: true -->
```OCaml
let remove_stutter l =
  (* YOUR CODE HERE *)
   raise (Failure "Not implemented")
```

<!-- grade_id: cell-aa3fb845581c6fce, points: 15, grade: true, solution: false -->
```OCaml
assert (remove_stutter [1;2;2;3;1;1;1;4;4;2;2] = [1; 2; 3; 1; 4; 2]) 
```

## Problem 7

Write a function

```ocaml
flatten : 'a list list -> 'a list
```

that flattens a list. For example, `flatten [[1;2];[3;4]] = [1;2;3;4]`.

<!-- grade_id: cell-1ba30bbe74fd3005, points: 0, grade: false, solution: true -->
```OCaml
let rec flatten l =
  (* YOUR CODE HERE *)
   raise (Failure "Not implemented")
```

<!-- grade_id: cell-8e986c3f5dbf5e89, points: 10, grade: true, solution: false -->
```OCaml
assert (flatten ([[1;2];[3;4]]) = [1;2;3;4])
```

## Problem 8

Write a function

```ocaml

type 'a tree = Leaf | Node of 'a tree * 'a * 'a tree

fold_inorder : ('a -> 'b -> 'a) -> 'a -> 'b tree -> 'a
```

That does a inorder fold of the tree. For example,

`fold_inorder (fun acc x -> acc @ [x]) [] (Node (Node (Leaf,1,Leaf), 2, Node (Leaf,3,Leaf))) = [1;2;3]`

<!-- grade_id: cell-7170388040d2ef4b, points: 0, grade: false, solution: true -->
```OCaml
type 'a tree = Leaf | Node of 'a tree * 'a * 'a tree

let rec fold_inorder f acc t =
  (* YOUR CODE HERE *)
   raise (Failure "Not implemented")
```

<!-- grade_id: cell-8bd22519c4af64bb, points: 15, grade: true, solution: false -->
```OCaml
assert (fold_inorder (fun acc x -> acc @ [x]) [] (Node (Node (Leaf,1,Leaf), 2, Node (Leaf,3,Leaf))) = [1;2;3]) 
```

## Problem 9

The usual recursive formulation of fibonacci function 

```ocaml
let rec fib n = 
  if n = 0 then 0
  else if n = 1 then 1
  else fib (n-1) + fib (n-2)
```

has exponential running time. It will take a long time to compute `fib 50`. You might have to interrupt the kernel if you did try to do `fib 50` in the notebook.

But we know that fibonacci number can be computed in linear time by remembering just the current `cur` and the previous `prev` fibonacci number. In this case, the next fibonacci number is computed as the sum of the current and the previous numbers. Then the program continues by setting `prev` to be `cur` and `cur` to be `cur + prev`.

Implement a tail recursive function `fib_tailrec` that uses this idea and computes the `nth` fibonacci number in linear time. 

```ocaml
fib_tailrec : int -> int
```

<!-- grade_id: cell-39286aaf6b703959, points: 0, grade: false, solution: true -->
```OCaml
let fib_tailrec n =
  (* YOUR CODE HERE *)
   raise (Failure "Not implemented")
```

<!-- grade_id: cell-57d3e7608b2286f5, points: 15, grade: true, solution: false -->
```OCaml
assert (fib_tailrec 50 = 12586269025)
```